```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

# Lab 0: Test Lab

A test of the different types of questions we can have.

---

## 1. Code


```python
# Run this cell to set up the notebook, but please don't change it.

import numpy as np
from datascience import *
%matplotlib inline
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

#### Part 1.1 (4 points) (1 pts)

Make an array called `numbers` containing the following numbers (in the given order)

1. 1
2. 2
3. 3

*Note:* Please make an **array** with `make_array`, and do not use a Python list (which behaves differently).

```python
numbers = ...
numbers
```

```python
grader.check("1.1-make_array")
```

---

## 2. Manually Graded Text Question


<!-- BEGIN QUESTION -->

**2.1 A Question.**
This is the question

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## 3. Manually Graded Question with some code


<!-- BEGIN QUESTION -->

**3.1 Another Question** This is the coding question

```python
x = ...
```

<!-- END QUESTION -->

## 4. Some datascience dude!


<!-- BEGIN QUESTION -->

#### 4.1 Make a Plot

```python
table = Table.read_table('temperatures.csv')
table.scatter(0,1,fit_line=True)
```

<!-- END QUESTION -->

## You're Done!

**Important submission information:** Be sure to run the tests and verify that they all pass, then choose **Save Notebook** from the **File** menu, then **run the final cell** and click the link to download the zip file. Then, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to the corresponding assignment. The name of this assignment is "Lab 0". **Be sure your work is saved before running the last cell!**

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(filtering=False)
```

Tests for 1.1-make_array (points: 4)
```
>>> len(numbers)
3
```
failure_message: Be sure to create an array with 3 elements.
