```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Homework 2: Arrays and Tables

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](http://data8.org/fa21/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings:**
- [Arrays](https://inferentialthinking.com/chapters/05/1/Arrays.html)
- [What is Data Science?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Deadline:**

This assignment is due Thursday, 9/9 at 11:59pm PT. Turn it in on Wednesday by 11:59pm PT for 5 extra credit points. Late work will not be accepted as per the [policies](http://data8.org/fa21/policies.html) page.

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the policies page to learn more about how to learn cooperatively.

You should start early so that you have time to get help if you're stuck. Office hours are held Monday-Friday. The schedule appears on [http://data8.org/fa21/office-hours.html](http://data8.org/fa21/office-hours.html).

## 1. Creating Arrays

```python
# Run this cell to set up the notebook, but please don't change it.

import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

### Part 1.1 (4 points)

Make an array called `numbers` containing the following numbers (in the given order)

1. -2
2. the floor of 12.6
3. 3
4. 5 to the power of the ceil of 5.3

*Hint:* `floor` and `ceil` are functions in the `math` module. Importing modules is covered in 2.1 of Lab 2!

*Note:* Python lists are different/behave differently than NumPy arrays. In Data 8, we use NumPy arrays, so please make an **array**, not a Python list.

```python
# Our solution involved one extra line of code before creating
# numbers.
...
numbers = ...
numbers
```

```python
grader.check("q1_1")
```

### Part 1.2 (4 points)

Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves".

```python
book_title_words = ...
book_title_words
```

```python
grader.check("q1_2")
```

### Part 1.3 (4 points)

Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **with** `a_string` inserted in between each string.

Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this one `without_commas`)

```python
with_commas = ...
without_commas = ...

# These lines are provided just to print out your answers.
print('with_commas:', with_commas)
print('without_commas:', without_commas)
```

```python
grader.check("q1_3")
```

## 2. Indexing Arrays

These exercises give you practice accessing individual elements of arrays with the `array.item(index)` method.  In Python, elements are accessed by its *index*; for example, the first element is the element at index 0. Indices must be **integers**.

***Note:* If you have previous coding experience, you may be familiar with bracket notation. DO NOT use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.**

Be sure to refer to the [Python Reference](http://data8.org/fa21/python-reference.html) on the website if you feel stuck!

### Part 2.1 (4 points)

The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers`.

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("q2_1")
```

### Part 2.2 (4 points)

The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.

*Hint:* Replace the `...` with strings or numbers. As a reminder, indices should be **integers**.

```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", blank_a, blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, 2, blank_d, 4),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("q2_2")
```

### Part 2.3 (4 points)

You'll sometimes want to find the **last** element of an array.  Suppose an array has 142 elements.  What is the index of its last element?

```python
index_of_last_element = ...
```

```python
grader.check("q2_3")
```

### Part 2.4 (4 points)

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

The cell below loads an array called `president_birth_years`.  Calling `tbl.column(...)` on a table returns an array of the column specified, in this case the `Birth Year` column of the `president_births` table. The last element in that array is the most recent among the birth years of all the deceased Presidents. Assign that year to `most_recent_birth_year`.

```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q2_4")
```

### Part 2.5 (4 points)

Finally, assign `min_of_birth_years` to the minimum of the first, sixteenth, and last birth years listed in `president_birth_years`.

```python
min_of_birth_years = min(president_birth_years.item(0), president_birth_years.item(15), most_recent_birth_year)
min_of_birth_years
```

```python
grader.check("q2_5")
```

**Question 3.** Suppose the surveyors started watching Old Faithful at the start of the first eruption. Assume that they watch until the end of the tenth eruption. For some of that time they will be watching eruptions, and for the rest of the time they will be waiting for Old Faithful to erupt. How many minutes will they spend waiting for eruptions?" **(4 Points)**

*Hint #1:* One way to approach this problem is to use the `take` or `where` method on the table `faithful_with_eruption_nums`. 

*Hint #2:* `first_nine_waiting_times` must be an array.

<!--
BEGIN QUESTION
name: q4_3
points:
 - 0
 - 4
-->

## You're Done!

You're done with Homework 2!  

**Important submission information:** Be sure to run the tests and verify that they all pass, then choose **Save and Checkpoint** from the **File** menu, then **run the final cell** and click the link to download the zip file. Then, go to [Gradescope](https://www.gradescope.com/courses/288788) and submit the zip file to the corresponding assignment. The name of this assignment is "Homework 02 Autograder". **It is your responsibility to make sure your work is saved before running the last cell.**

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: 4)
```
>>> len(numbers)
4
```
failure_message: Be sure to create an array with four elements.


Tests for q1_2 (points: 4)
```
>>> import numpy as np
>>> type(book_title_words) == np.ndarray
True
```
failure_message: Be sure to use make_array to create an array.
```
>>> len(book_title_words) == 3
True
```
```
>>> book_title_words.item(0)
'Eats'
```
```
>>> book_title_words.item(1)
'Shoots'
```
```
>>> book_title_words.item(2)
'and Leaves'
```


Tests for q1_3 (points: 4)
```
>>> with_commas
'Eats, Shoots, and Leaves'
```
```
>>> without_commas
'Eats Shoots and Leaves'
```


Tests for q2_1 (points: 4)
```
>>> third_element != -10
True
```
failure_message: Remember that indices start at 0...


Tests for q2_2 (points: 4)
```
>>> elements_of_some_numbers.column(0).item(2)
'third'
```
```
>>> elements_of_some_numbers.column(0).item(3)
'fourth'
```
```
>>> elements_of_some_numbers.column(1).item(0)
0
```
```
>>> elements_of_some_numbers.column(1).item(3)
3
```


Tests for q2_3 (points: 4)
```
>>> index_of_last_element
141
```


Tests for q2_4 (points: 4)
```
>>> most_recent_birth_year
1917
```


Tests for q2_5 (points: 4)
```
>>> min_of_birth_years > 0
True
```
