```python
# Initialize Otter
import otter
grader = otter.Notebook("pre06.ipynb")
```

# Prelab 6: Models and Hypothesis Testing

**Helpful Resources:**
- [Python Reference](http://www.cs.williams.edu/~cs104/auto/python-library-ref.html): Cheat sheet of helpful library methods.

**Readings:**
* [Ch 10. Sampling and Empirical Distributions](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Ch 11.1. Assessing a Model](https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html)


Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.  For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this prelab and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

```python
# Run this cell to set up the notebook.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines make plots look nice and hide some messy Python warnings.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
warnings.simplefilter('ignore', np.VisibleDeprecationWarning)

# For special interactions
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
```

## 1. The Therapeutic Touch (40 pts)




The Therapeutic Touch (TT) is the idea that everyone can feel the Human Energy Field (HEF) around individuals.  Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy." 

TT was a popular technique used throughout the 20th century that was toted as a great way to bring balance to a person's health. Certain practitioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals.

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT.

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly interact with a person's HEF. She later went on to publish her work in TT, becoming the youngest person to have a research paper published in a peer reviewed medical journal.

### Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner and ask them to extend their hands through a screen (which they can't see through). Emily would be on the other side and would flip a fair coin. Depending on how the coin landed, she would put out either her left hand or her right hand. The TT practitioner would then have to answer which hand Emily put out. If a pracitioner could truly interact with a person's HEF, it would be expected that they answered correctly.

Overall, through 210 samples, the practitioner picked the correct hand 44% of the time. 

Emily's main goal here was to test whether or not the TT practicioners' guesses were random, like the flip of a coin. In most medical experiments, this is the norm. We want to test whether or not the treatment has an effect, *not* whether or not the treatment actually works. 

We will now begin to formulate this experiment in terms of the terminology we learned in this course:
* Emily's model is that the TT practitioners have a 50% chance of choosing the correct hand, like flipping a fair coin.
* Emily is trying to debunk the alternative model that the TT practitioners have some chance other than 50% of choosing the correct hand. 

#### Part 1.1 (5 pts)


Remember that the practitioner got the correct answer 44% (0.44) of the time. According to Emily's model, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is between 0 and 1. 

```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("q1.1")
```

#### Part 1.2 (5 pts)


The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 

We usually use a statistic to help determine which model the evidence points towards. What is a statistic that we can use to compare outcomes under Emily’s model to what was observed? Assign `valid_stat` to an array of integer(s) representing test statistics that Emily can use: 

1. The difference between the expected percent correct and the actual percent correct
2. The absolute difference between the expected percent correct and the actual percent correct
3. The sum of the expected percent correct and the actual percent correct



```python
valid_stat = make_array(2)
valid_stat
```

```python
grader.check("q1.2")
```

If you had trouble coming up with the correct answer, keep in mind that we are interested in exploring whether or not the treatment has an effect, not whether or not the treatment actually works. This is because the alternative model states that practitioners are not answering randomly, but not whether their accuracy is better or worse than if they were guessing at random. Thus, we want to take the absolute difference between the expected percent correct and actual percent correct to account for the alternative model being either better or worse than guessing at random. 

#### Part 1.3 (5 pts)


 Define the function `statistic` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in Question 3. Assume that the argument takes in proportions, but  return your answer as a percentage. 

*Hint:* Remember we are asking for a **percentage**, not a proportion. 


```python
def statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("q1.3")
```

#### Part 1.4 (5 pts)


Use your newly defined function to calculate the observed statistic from Emily's experiment. 

```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q1.4")
```

### Using `sample_proportions`

**Now we will test whether this observed statistic consistent with what we might see under Emily’s model.**

In order to answer this question, we must simulate the experiment as though Emily's model was correct, and calculate our statistic for every simulation.

`sample_proportions` can be used to randomly sample from multiple categories when you know the proportion of data points that are expected to fall in each category. `sample_proportions` takes two arguments: the sample size and an array that contains the distribution of categories in the population (should sum to 1).

Consider flipping a fair coin, where the two outcomes (coin lands heads and coin lands tails) occur with an equal chance. We expect that half of all coin flips will land heads, and half of all coin flips will land tails.

Run the following cell to see the simulation of 10 flips of a fair coin. Let the first item of `coin_proportions` be the proportion of heads and the second item of `coin_proportions` be the proportion of tails.

```python
coin_proportions = make_array(0.5, 0.5) 
ten_flips = sample_proportions(10, coin_proportions)
ten_flips
```

`sample_proportions` returns an array that is the same length as the proportion array that is passed through. It contains the proportion of each category that appears in the sample. 

In our example, the first item of `ten_flips` is the simulated proportion of heads and the second item of `ten_flips` is the simulated proportion of tails.

```python
simluated_proportion_heads = ten_flips.item(0)
simluated_proportion_tails = ten_flips.item(1)

print("In our simluation, " + str(simluated_proportion_heads) + " of flips were heads and " \
      + str(simluated_proportion_tails) + " of flips were tails.")
```

#### Part 1.5 (5 pts)


 To begin simulating, we should start by creating a representation of Emily's model to use for our simulation. This will be an array with two items in it. The first item should be the proportion of times, assuming that Emily’s model was correct, a TT practictioner picks the correct hand. The second item should be the proportion of times, under the same assumption, that the TT practitioner picks the incorrect hand. Assign `model_proportions` to this array. 

After this, we can simulate 210 hand choices, as Emily evaluated in real life, and find a single statistic to summarize this instance of the simulation. Use the `sample_proportions` function and assign the proportion of correct hand choices (out of 210) to `simulation_proportion_correct`. Lastly, use your statistic function to assign `one_statistic`  to the value of the statistic for this one simulation.

*Hint:* `sample_proportions` usage can be found [here](http://data8.org/su19/python-reference.html).


```python
model_proportions = ...
simulation_proportion_correct = ...
one_statistic = ...
one_statistic
```

```python
grader.check("q1.5")
```

#### Part 1.6 (5 pts)


 Let's now see what the distribution of statistics is actually like under Emily's model. 

Define the function `simulation_and_statistic` to take in the `model_proportions` array and the expected proportion of times a TT practitioner would guess a hand correctly under Emily's model. The function should simulate Emily running through the experiment 210 times and return the statistic of this one simulation. 

*Hint:* This should follow the same pattern as the code you did in the previous problem.  


```python
def simulation_and_statistic(model_proportions, expected_proportion_correct):
    '''Simulates 210 TT hand choices under Emily’s model. 
    Returns one statistic from the simulation.'''
    ...
```

#### Part 1.7 (5 pts)


Using this function, assign `simulated_statistics` to an array of 1000 statistics that you calculated under the assumption that Emily's model was true.


```python
num_repetitions = 1000

simulated_statistics = ...

...
...

```

```python
grader.check("q1.7")
```

Let's view the distribution of the simulated statistics under Emily's model, and visually compare where the observed statistic lies relative to the simulated statistics.

```python
t = Table().with_column('Simulated Statistics', simulated_statistics)
t.hist()
plt.scatter(observed_statistic, 0, color='red', s=30);
```

We can make a visual argument as to whether we believe the observed statistic is consistent with Emily’s model. Here, since larger values of the test statistic suggest the alternative model (where the chance of guessing the correct hand is something other than 50%), we can formalize our analysis by finding what proportion of simulated statistics were as large or larger than our observed test statistic (the area at or to the right of the observed test statistic). If this area is small enough, we’ll declare that the observed data are inconsistent with our simulated model.

#### Part 1.8 (5 pts)


Calculate the proportion of simulated statistics greater than or equal to the observed statistic. 


```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("q1.8")
```

By convention, we often compare the proportion we just calculated to 0.05. If the proportion of simulated statistics greater than or equal to the observed statistic is sufficiently small (less than or equal to 0.05), then this is evidence against Emily's model. Otherwise, we don’t have any reason to doubt Emily’s model. 

This should help you make your own conclusions about Emily Rosa's experiment. 

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners hit back and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want!

Think about the following questions as to what conclusions you can draw from this study: 

1. Is the data more consistent with Emily' model (practioners were randomly guessing)?
2. What does this mean in terms of Emily's experiment? Do the TT practitioners' answers follow an even chance model or is there something else at play? 

## 2. You're Done!


**Important submission information:** Follow these steps to submit your work:
* Run the tests and verify that they pass as you expect. 
* Choose **Save Notebook** from the **File** menu.
* **Run the final cell** and click the link below to download the zip file. 

Once you have downloaded that file, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to the corresponding assignment. The name of this assignment is "Prelab 4 Autograder". **Be sure your work is saved before running the last cell!**

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.1 (points: 5)
```
>>> 0 <= expected_proportion_correct <= 1
True
```
```
>>> expected_proportion_correct == 0.5
True
```


Tests for q1.2 (points: 5)
```
>>> type(valid_stat) == np.ndarray
True
```
```
>>> int(sum(valid_stat)) == 2
True
```


Tests for q1.3 (points: 5)
```
>>> int(round(statistic(.5,.5) + statistic(.4,.1),1)) == 30
True
```
```
>>> int(statistic(.4,.1) - statistic(.1,.4)) == 0
True
```


Tests for q1.4 (points: 5)
```
>>> int(round(observed_statistic,2)) == 6
True
```


Tests for q1.5 (points: 5)
```
>>> len(model_proportions) % 2 == 0
True
```
```
>>> len(np.unique(model_proportions)) == 1
True
```
```
>>> sum(model_proportions) == 1
True
```
```
>>> type(simulation_proportion_correct) == float
True
```
```
>>> 0 <= one_statistic <= 25
True
```


Tests for q1.7 (points: 5)
```
>>> len(simulated_statistics) == 1000
True
```
```
>>> np.all(simulated_statistics <= 30)
True
```
```
>>> np.all(simulated_statistics >= 0)
True
```
```
>>> 0 <= simulation_and_statistic(model_proportions, expected_proportion_correct) <= 25
True
```


Tests for q1.8 (points: 5)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
```
>>> proportion_greater_or_equal*1000 == np.count_nonzero(simulated_statistics >= observed_statistic)
True
```
