```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# Lab 1: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

## 1. Scary Arithmetic (5 pts)



<!-- BEGIN QUESTION -->

#### Part 1.1 (5 pts)


An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day to Labor Day."

Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day to Labor Day? Please explain your answer.

**Notes:** You can assume that "over 25%" means only slightly over. Had it been much over, say closer to 30%, then the marketers would have said so. Also, Memorial Day is observed on the last Monday of May and Labor Day is observed on the first Monday of September.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Characters in Little Women (10 pts)



In lecture, we counted the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://www.inferentialthinking.com/chapters/01/3/1/literary-characters). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

#### Part 2.1 (5 pts)


Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q1` below.

1. 250
2. 390
3. 440
4. 32,000
5. 40,000

```python
characters_q1 = ...
```

```python
grader.check("q2.1")
```

The test above checks that your answers are in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully in each multiple-choice answer cell.

#### Part  2.2 (5 pts)


Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q2` below.
1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods

```python
characters_q2 = ...
```

```python
grader.check("q2.2")
```

Again, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, read [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) of the textbook.

## 3. Names and Assignment Statements (15 pts)



#### Part  3.1 (5 pts)


When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer.

1. Python is smart and already knows `4 = 2 + 2`.

2. `4` is already a defined number, and it doesn't make sense to make a number be a name for something else. In Python, "`x = 2 + 2`" means "assign `x` as the name for the value of `2 + 2`."

3. It should be `2 + 2 = 4`.

```python
names_q1 = ...
```

```python
grader.check("q3.1")
```

#### Part  3.2 (5 pts)


When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. Python cannot interpret the name `plus`, as it has not been defined.

<!--
BEGIN QUESTION
name: q3_2
manual: False
points:
 - 0
 - 4
-->

```python
names_q2 = ...
```

```python
grader.check("q3.2")
```

#### Part 3.3 (5 pts)


When you run the following cell, Python will, yet again, produce another cryptic error message.

```python
x = print(5)
y = x + 2
```

Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_q3` below to indicate your answer.

1. Python doesn't want `y` to be assigned.

2. The `print` operation is meant for displaying values to the programmer, not for assigning values!

3. Python can’t do addition between one name and one number. It has to be 2 numbers or 2 predefined names.

<!--
BEGIN QUESTION
name: q3_3
manual: false
points:
 - 0
 - 4
-->

```python
names_q3 = ...
```

```python
grader.check("q3.3")
```

## 4. Differences between Majors (20 pts)



Adapted from information found on the Williams website, the table below displays the average number of degree recipients in three majors for the years 2008-2012 and 2011-2222.

| Major                              | 2008-2012    | 2018-2021   |
|------------------------------------|--------------|-------------|
| Comparative Literature             |  5           | 11          |
| Psychology                         | 62           | 45          |
| Mathematics                        | 53           | 61          |


#### Part 4.1 (5 pts)


Suppose you want to find the **biggest** absolute difference between the numbers of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 5 from 11) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers.

```python
biggest_change = ...
biggest_change
```

```python
grader.check("q4.1")
```

#### Part 4.2 (5 pts)


Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1. Comparative Literature
2. Psychology
3. Mathematics

Choose the number that corresponds to the major with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major.

```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("q4.2")
```

#### Part 4.3 (5 pts)


For each major, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2012}} * 100}$ 

Fill in the code below such that `gws_relative_change`, `linguistics_relative_change` and `rhetoric_relative_change` are assigned to the relative changes for their respective majors.

```python
complit_relative_change = (abs(...) / 5) * 100

psych_relative_change = ...
math_relative_change = ...
complit_relative_change, psych_relative_change, math_relative_change
```

```python
grader.check("q4.3")
```

#### Part 4.4 (5 pts)


Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1. Comparative Literature
2. Psychology
3. Mathematics

Choose the number that corresponds to the major with the biggest relative change. 

```python
# Assign biggest_rel_change_major to the number corresponding to the major with the biggest relative change.
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("q4.4")
```

## 5. Nearsightedness Study (20 pts)



Myopia, or nearsightedness, results from a number of genetic and environmental factors. In 1999, Quinn et al studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

<!-- BEGIN QUESTION -->

#### Part 5.1 (5 pts)


 The data were gathered by the following procedure, reported in the study. "Between January and June 1998, parents of children aged 2-16 years [...] that were seen as outpatients in a university pediatric ophthalmology clinic completed a questionnaire on the child’s light exposure both at present and before the age of 2 years." Was this study observational, or was it a controlled experiment? Explain. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part 5.2 (5 pts)


 The study found that of the children who slept with a room light on before the age of 2, 55% were myopic. Of the children who slept with a night light on before the age of 2, 34% were myopic. Of the children who slept in the dark before the age of 2, 10% were myopic. The study concluded that, "The prevalence of myopia [...] during childhood was strongly associated with ambient light exposure during sleep at night in the first two years after birth."

Do the data support this statement? You may interpret "strongly" in any reasonable qualitative way.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part 5.3 (5 pts)


 On May 13, 1999, CNN reported the results of this study under the headline, "Night light may lead to nearsightedness." Does the conclusion of the study claim that night light causes nearsightedness?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part 5.4 (5 pts)


 The final paragraph of the CNN report said that "several eye specialists" had pointed out that the study should have accounted for heredity.

Myopia is passed down from parents to children. Myopic parents are more likely to have myopic children, and may also be more likely to leave lights on habitually (since the parents have poor vision). In what way does the knowledge of this possible genetic link affect how we interpret the data from the study? 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

Myopia, or nearsightedness, results from a number of genetic and environmental factors. In 1999, Quinn et al studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

## 6. Studying the Survivors (5 pts)



#### Part 6.1 (5 pts)


The Reverend Henry Whitehead was skeptical of John Snow’s conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong.  (The history of the event is detailed [here](http://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group?

1) If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.

2) Survivors could provide additional information about what else could have caused the cholera, potentially unearthing another cause.

3) Through considering the survivors, Whitehead could have identified a cure for cholera.

```python
# Assign survivor_answer to 1, 2, or 3
survivor_answer = ...
```

```python
grader.check("q6.1")
```

**Note:** Whitehead ended up finding further proof that the Broad Street pump played the central role in spreading the disease to the people who lived near it. Eventually, he became one of Snow’s greatest defenders.

## 7. Job Opportunities & Education in Rural India (15 pts)



#### Part 7.1 (5 pts)


Which statement best describes the *treatment* and *control* groups for this study? Assign either 1, 2, or 3 to the name `jobs_q1` below.

1. The treatment group was the 80 villages visited by recruiters, and the control group was the other 80 villages with no intervention.

2. The treatment group was the 160 villages selected, and the control group was the rest of the villages outside the capital of India.

3. There is no clear notion of *treatment* and *control* group in this study.

```python
jobs_q1 = ...
```

```python
grader.check("q7.1")
```

#### Part 7.2 (5 pts)


 Was this an observational study or a randomized controlled experiment? Assign either 1, 2, or 3 to the name `jobs_q2` below.

1. This was an observational study.

2. This was a randomized controlled experiment.  

3. This was a randomized observational study.

```python
jobs_q2 = ...
```

```python
grader.check("q7.2")
```

<!-- BEGIN QUESTION -->

#### Part 7.3 (5 pts)


 The study reported, "Girls aged 5-15 in villages that received the recruiting services were 3 to 5 percentage points more likely to be in school and experienced an increase in Body Mass Index, reflecting greater nutrition and/or medical care. However, there was no net gain in height. For boys, there was no change in any of these measures." Why do you think the author points out the lack of change in the boys?

*Hint:* Remember the original hypothesis. The author believes that educating women in job opportunities will cause families to invest more in the women’s well-being.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 8. Welcome Survey


**Question 1:** Once you have submitted, please also complete the welcome survey in order to receive credit for homework 1.

Welcome survey is here: https://forms.gle/pUbSa7AW1xVvnicYA

Assign `survey` to the secret string given at the end of the welcome survey:

<!--
BEGIN QUESTION
name: q8_1
points:
 - 1
-->

```python
survey = ...
```

```python
# grader.check("q8_1")
```

## You're done!

**Important submission information:** Be sure to run the tests and verify that they all pass, then choose **Save Notebook** from the **File** menu, then **run the final cell** and click the link to download the zip file. Then, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to the corresponding assignment. The name of this assignment is "Lab 01 Autograder". **It is your responsibility to make sure your work is saved before running the last cell.**

Once you have submitted, your Gradescope assignment should look something like the following image if you have passed all tests.

NOTE: *This is an image of a generic Gradescope submission result --- it does not included the same test numbers as this assignment's.*

<img src="gradescope.png">

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q2.1 (points: None)
```
>>> 1 <= characters_q1 <= 5
```


Tests for q2.2 (points: None)
```
>>> 1 <= characters_q2 <= 5
```


Tests for q3.1 (points: None)
```
>>> 1 <= names_q1 <= 3
```


Tests for q3.2 (points: None)
```
>>> 1 <= names_q2 <= 4
```


Tests for q3.3 (points: None)
```
>>> 1 <= names_q3 <= 4
```


Tests for q4.1 (points: None)
```
>>> 0 <= biggest_change
```


Tests for q4.2 (points: None)
```
>>> 1 <= smallest_change_major <= 3
```
```
>>> smallest_change_major == 1
```


Tests for q4.3 (points: None)
```
>>> 0 <= complit_relative_change <= 200
```
```
>>> 0 <= psych_relative_change <= 200
```
```
>>> 0 <= math_relative_change <= 200
```
```
>>> #HIDDEN
>>> np.isclose(complit_relative_change, 120.0)
```
```
>>> #HIDDEN
>>> np.isclose(psych_relative_change, 27.419354838709676)
```
```
>>> #HIDDEN
>>> np.isclose(math_relative_change, 15.09433962264151)
```


Tests for q4.4 (points: None)
```
>>> 1 <= biggest_rel_change_major <= 3
```


Tests for q6.1 (points: None)
```
>>> 1 <= survivor_answer <= 3
```


Tests for q7.1 (points: None)
```
>>> # TEST
>>> 1 <= jobs_q1 <= 3
```


Tests for q7.2 (points: None)
```
>>> # TEST
>>> 1 <= jobs_q2 <= 3
```
