```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

# Lab 5: Applying Functions and Iteration

**Helpful Resources:**
- [Python Reference](http://www.cs.williams.edu/~cs104/python-library-ref.html): Cheat sheet of helpful library methods.

**Recommended Readings**: 
* [Ch 8.1. Applying Functions](https://www.inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Ch 9.1. Conditionals](https://www.inferentialthinking.com/chapters/09/1/Conditional_Statements.html)
* [Ch 9.2. Iteration](https://www.inferentialthinking.com/chapters/09/2/Iteration.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.  For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

## 1. Eph Field Hockey (20 pts)



```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines make plots look nice and hide some messy Python warnings.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
warnings.simplefilter('ignore', np.VisibleDeprecationWarning)
```

Field Hockey games are divided into four quarters, with an overtime period if a game is tied after four quarters. The number of points the Ephs scored in each quarter, and the number of points their opponent scored in each quarter during their 2021 season are stored in a table called `williams-fh-2021.csv`.

```python
# Just run this cell
# Read in the cal_fb csv file
games = Table().read_table("williams-fh-2021.csv")
games.show()
```

Let's start by finding the total points each team scored in a game.

#### Part 1.1 (5 pts)


 Write a function called `sum_scores`.  It should take four arguments, where each argument is the team's score for that quarter. It should return the team's total score for that game. 

*Tip:* Don't overthink this question!




```python
def sum_scores(...):
     '''Returns the total score calculated by adding up the score of each quarter'''
     ...
"""; # END PROMPOT

sum_scores(14, 7, 3, 0, 5) # Example of how to call your function
```

```python
grader.check("q1.1")
```

#### Part 1.2 (5 pts)


Create a new table `final_scores` with three columns in this *specific* order: `Opponent`, `Ephs Score`, `Opponent Score`. You will have to create the `Ephs Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem. 

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the python reference for syntax.

*Tip:* If you’re running into issues creating `final_scores`, check that `ephs_scores` and `opponent_scores` output what you want. 

```python
ephs_scores = ...
opponent_scores = ...
final_scores = ...
final_scores
```

```python
grader.check("q1.2")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `n`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
games.row(9)
```

```python
# Just run this cell
games.row(9).item("Ephs 4Q")
```

#### Part 1.3 (5 pts)


We want to see for a particular game whether or Williams won or lost. Write a function called `did_ephs_win`.  It should take one argument: a **row object** from the `final_scores` table. The `did_ephs_win` function should return either `True` if Williams's score was greater than the Opponent's score, and `False` otherwise.

*Note 1*: "Row object" means a row from the table that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_scores.row(row)` inside of the function.

*Note 2*: If you're still confused by row objects, try printing out `final_scores.row(1)` in a new cell to visually see what it looks like! This piece of code is pulling out the row object located at index 1 of the `final_scores` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!

```python
def did_ephs_win(row):
    ...
```

```python
grader.check("q1.3")
```

#### Part 1.4 (5 pts)


We'll use your functions to compute the outcome of every game in the season. Using the `final_scores` table:
1. Assign `ephs_won` to an array of `True` and `False` values that correspond to whether or not Williams won.
2. Add the `ephs_won` array to the `final_scores` table in a column named `Won`, and assign this to `final_scores_with_results`.
3. Then, respectively assign the number of wins and losses Williams had to `eph_wins` and `eph_losses`.

*Hint 1*: `True` and `False` are **not** strings " ". What data type are they?

*Hint 2*: `tbl.num_rows` might be helpful too.

*Hint 3*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`.




```python
results = ...
final_scores_with_results = ...
eph_wins = ...
eph_losses = ...

# Don't delete or edit the following line:
print(f"In the 2019 Season, Williams won {eph_wins} games and lost {eph_losses} games. Go Ephs!")
```

```python
grader.check("q1.4")
```

You're done with the required section of Lab 5! Continue on to the optional section for some more practice with iterations and for loops! 

Make sure to **run the submit cell located at the bottom of this notebook.**

## 2. Unrolling Loops (Optional) -- This question is awful


**The rest of this homework is optional. Do it for your own practice, but it will not be incorporated into the final grading!**

"Unrolling" a `for` loop means to manually write out all the code that it executes.  The result is code that does the same thing as the loop, but without the structure of the loop.  For example, for the following loop:

    for num in np.arange(3):
        print("The number is", num)

The unrolled version would look like this:

    print("The number is", 0)
    print("The number is", 1)
    print("The number is", 2)


Unrolling a `for` loop is a great way to understand what the loop is doing during each step. In this exercise, you'll practice unrolling a `for` loop.


In the question below, write code that does the same thing as the given code, but with any `for` loops unrolled.  It's a good idea to run both your answer and the original code to verify that they do the same thing.  (Of course, if the code does something random, you'll get a different random outcome than the original code!)

**Optional Question 1.** Unroll the code below.


```python
for joke_iteration in np.arange(3):
    print("Knock, knock.")
    print("Who's there?")
    print("Banana.")
    print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
```

```python
...  # SOLUTION

```

You're done with Lab 5!  

**Important submission information:** Follow these steps to submit your work:
* Run the tests and verify that they pass as you expect. 
* Choose **Save Notebook** from the **File** menu.
* **Run the final cell** and click the link below to download the zip file. 

Once you have downloaded that file, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to the corresponding assignment. The name of this assignment is "Lab 5". **Be sure your work is saved before running the last cell!**

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1.1 (points: None)
```
>>> sum_scores(2, 3, 6, 1, 0) == 12
True
```
```
>>> sum_scores(-2, 3, 5, -11, 1) == -4
True
```


Tests for q1.2 (points: None)
```
>>> final_scores.num_columns == 3
True
```
```
>>> final_scores.labels == ('Opponent', 'Ephs Score', 'Opponent Score')
True
```


Tests for q1.3 (points: None)
```
>>> did_ephs_win(final_scores.row(0)) == False
True
```
```
>>> did_ephs_win(final_scores.row(1)) == True
True
```


Tests for q1.4 (points: None)
```
>>> 0 <= eph_wins <= 10
True
```
```
>>> 0 <= eph_losses <= 10
True
```
