```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

# Lab 3: Table Manipulation and Visualization

**Helpful Resources:**
- [Python Reference](http://www.cs.williams.edu/~cs104/python-library-ref.html): Cheat sheet of helpful library methods.

**Recommended Reading**: 
* [Ch 7. Visualization](https://inferentialthinking.com/chapters/07/Visualization.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## 1. Unemployment (35 pts)



The Federal Reserve Bank of St. Louis publishes data about jobs in the US.  Below, we've loaded data on unemployment in the United States. There are many ways of defining unemployment, and our dataset includes two notions of the unemployment rate:

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI-PTER.  (Economists are great at marketing.)

The source of the data is [here](https://fred.stlouisfed.org/categories/33509).

#### Part 1.1 (5 pts)


 The data are in a CSV file called `unemployment.csv`.  Load that file into a table called `unemployment`. 



```python
unemployment = ...
unemployment
```

```python
grader.check("q1.1")
```

#### Part 1.2 (5 pts)


 Sort the data in descending order by NEI, naming the sorted table `by_nei`.  Create another table called `by_nei_pter` that's sorted in descending order by NEI-PTER instead. 



```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("q1.2")
```

#### Part 1.3 (5 pts)


 Use `take` to make a table containing the data for the 10 quarters when NEI was greatest.  Call that table `greatest_nei`.

`greatest_nei` should be sorted in descending order of `NEI`. Note that each row of `unemployment` represents a quarter. 

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("q1.3")
```

#### Part 1.4 (5 pts)


 It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI-PTER is the percentage of people who are unemployed (included in the NEI) plus the percentage of people who are PTER.  Compute an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.) 

*Note:* Use the original `unemployment` table for this.



```python
pter = ...
pter
```

```python
grader.check("q1.4")
```

#### Part 1.5 (5 pts)


 Add `pter` as a column to `unemployment` (name the column `PTER`) and sort the resulting table by that column in descending order.  Call the resulting table `by_pter`.

Try to do this with a single line of code, if you can. 



```python
by_pter = ...
by_pter
```

```python
grader.check("q1.5")
```

<!-- BEGIN QUESTION -->

#### Part 1.6 (5 pts)


 

Create a line plot of PTER over time. 

To do this, create a new table called `pter_over_time` that augments the `year` array and the `pter` array to the `unemployment` table. Label these columns `Year` and `PTER`. Then, generate a line plot using one of the table methods you've learned in class.

The order of the columns matter for our correctness tests, so be sure `Year` comes before `PTER`. 

*Clarification: When constructing `pter_over_time`, do not just add the `year` column to the `by_pter` table. Please follow the directions in the question above.*



```python
year = 1994 + np.arange(by_pter.num_rows)/4
pter_over_time = ...
...

plots.ylim(0,2); # Do not change this line -- sets bounds on y-axis
```

```python
grader.check("q1.6")
```

<!-- END QUESTION -->

#### Part 1.7 (5 pts)


 Were PTER rates high during the Great Recession (that is to say, were PTER rates particularly high in the years 2008 through 2011)? Assign `highPTER` to `True` if you think PTER rates were high in this period, or `False` if you think they weren't. 



```python
highPTER = ...
```

```python
grader.check("q1.7")
```

## 2. Birth Rates (30 pts)



The following table gives census-based population estimates for each state on both July 1, 2015 and July 1, 2016. The last four columns describe the components of the estimated change in population during this time interval. **For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.**

The data was taken from [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2016/national/totals/nst-est2016-alldata.csv).

If you want to read more about the different column descriptions, click [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2015/national/totals/nst-est2015-alldata.pdf)!

The raw data is a bit messy - run the cell below to clean the table and make it easier to work with.

```python
# Don't change this cell; just run it.
pop = Table.read_table('nst-est2016-alldata.csv').where('SUMLEV', 40).select([1, 4, 12, 13, 27, 34, 62, 69])
pop = pop.relabeled('POPESTIMATE2015', '2015').relabeled('POPESTIMATE2016', '2016')
pop = pop.relabeled('BIRTHS2016', 'BIRTHS').relabeled('DEATHS2016', 'DEATHS')
pop = pop.relabeled('NETMIG2016', 'MIGRATION').relabeled('RESIDUAL2016', 'OTHER')
pop = pop.with_columns("REGION", np.array([int(region) if region != "X" else 0 for region in pop.column("REGION")]))
pop.set_format([2, 3, 4, 5, 6, 7], NumberFormatter(decimals=0)).show(5)
```

#### Part 2.1 (5 pts)


 Assign `us_birth_rate` to the total US annual birth rate during this time interval. The annual birth rate for a year-long period is the total number of births in that period as a proportion of the population size at the start of the time period. 

*Hint:* Which year corresponds to the start of the time period?



```python
us_birth_rate = ...
us_birth_rate
```

```python
grader.check("q2.1")
```

#### Part 2.2 (5 pts)


 Assign `movers` to the states for which the **absolute value** of the **annual rate of migration** was higher than 1%. The annual rate of migration for a year-long period is the net number of migrations (in and out) as a proportion of the population size at the start of the period. The `MIGRATION` column contains estimated annual net migration counts by state. 



```python
migration_rates = ...
movers = ...
movers
```

```python
grader.check("q2.2")
```

#### Part 2.3 (5 pts)


 Assign `west_births` to the total number of births that occurred in region 4 (the Western US). 

*Hint:* Make sure you double check the type of the values in the `REGION` column, and appropriately filter (i.e. the types must match!).



```python
west_births = ...
west_births
```

```python
grader.check("q2.3")
```

#### Part 2.4 (5 pts)


 

In the next question, you will be creating a visualization to understand the relationship between birth and death rates. The annual death rate for a year-long period is the total number of deaths in that period as a proportion of the population size at the start of the time period.

What visualization is most appropriate to see if there is an association between birth and death rates during a given time interval?

1. Line Graph
<br>
2. Bar Chart
<br>
3. Scatter Plot

Assign `visualization` below to the number corresponding to the correct visualization. 



```python
visualization = ...
```

```python
grader.check("q2.4")
```

<!-- BEGIN QUESTION -->

#### Part 2.5 (5 pts)


 In the code cell below, create a visualization that will help us determine if there is an association between birth rate and death rate during this time interval. It may be helpful to create an intermediate table here. 

Things to consider:
- What type of chart will help us illustrate an association between 2 variables?
- How can you manipulate a certain table to help generate your chart?

```python
# In this cell, use birth_props and death_props to generate your chart
birth_rates = pop.column('BIRTHS') / pop.column('2015')
death_rates = pop.column('DEATHS') / pop.column('2015')
...
```

<!-- END QUESTION -->

#### Part 2.6 (5 pts)


True or False: There is an association between birth rate and death rate during this time interval. 

Assign `assoc` to `True` or `False` in the cell below. 

```python
assoc = ...
```

```python
grader.check("q2.6")
```

## 3. Uber (30 pts)



Below we load tables containing 200,000 weekday Uber rides in the Manila, Philippines, and Boston, Massachusetts metropolitan areas from the [Uber Movement](https://movement.uber.com) project. The `sourceid` and `dstid` columns contain codes corresponding to start and end locations of each ride. The `hod` column contains codes corresponding to the hour of the day the ride took place. The `ride time` column contains the length of the ride, in minutes.

```python
boston = Table.read_table("boston.csv")
manila = Table.read_table("manila.csv")
print("Boston Table")
boston.show(4)
print("Manila Table")
manila.show(4)
```

<!-- BEGIN QUESTION -->

#### Part 3.1 (5 pts)


 Produce a histogram that visualizes the distributions of all ride times in Boston using the given bins. 



```python
equal_bins = np.arange(0, 120, 5)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part 3.2 (5 pts)


 Now, produce a histogram that visualizes the distribution of all ride times in Manila using the given bins. 



```python
equal_bins = np.arange(0, 120, 5)
...

# Don't delete the following line!
plots.ylim(0, 0.05);
```

<!-- END QUESTION -->

#### Part 3.3 (5 pts)


 Let's take a closer look at the y-axis label. Assign `unit_meaning` to an integer (1, 2, 3) that corresponds to the "unit" in "Percent per unit". 

1. minutes  
2. ride time  
3. seconds



```python
unit_meaning = ...
unit_meaning
```

```python
grader.check("q3.3")
```

#### Part 3.4 (5 pts)


 Assign `boston_under_15` and `manila_under_15` to the percentage of rides that are less than 15 minutes in their respective metropolitan areas. Use the height variables that we have defined for you below in order to compute the percentages. Your solution should only use height variables, numbers, and mathematical operations. You should **not** access the tables `boston` and `manila` in any way. 

```python
boston_under_5_height = 1.2
manila_under_5_height = 0.6
boston_5_to_under_10_height = 3.2
manila_5_to_under_10_height = 1.4
boston_10_to_under_15_height = 4.9
manila_10_to_under_15_height = 2.2

boston_under_15 = ...
manila_under_15 = ...

boston_under_15, manila_under_15
```

```python
grader.check("q3.4")
```

#### Part 3.5 (5 pts)


 Let's take a closer look at the distribution of ride times in Boston. Assign `boston_median_bin` to an integer (1, 2, 3, 4, or 5) that corresponds to the bin that contains the median time. 

1. 0-5 minutes  
2. 5-10 minutes  
3. 10-15 minutes  
4. 15-20 minutes  
5. 20-25 minutes  

*Hint:* The median of a sorted list has half of the list elements to its left, and half to its right.

```python
boston_median_bin = ...
boston_median_bin
```

```python
grader.check("q3.5")
```

<!-- BEGIN QUESTION -->

#### Part 3.6 (5 pts)


Identify one difference between the histograms, in terms of actual histogram properties. Based on the following two readings, why are the distributions for Boston and Manila different? 

- [Boston reading](https://www.climatestotravel.com/climate/united-states/boston)
- [Manila reading](https://newsinfo.inquirer.net/1383983/metro-manila-traffic-2nd-worst-in-the-world-report)

*Hint:* Try thinking about external factors of the two cities that may be causing the difference!



_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Histograms (20 pts)



Consider the following scatter plot: 

![](scatter.png)

The axes of the plot represent values of two variables: $x$ and $y$. 

Suppose we have a table called `t` that has two columns in it:

- `x`: a column containing the x-values of the points in the scatter plot
- `y`: a column containing the y-values of the points in the scatter plot

Below, you are given three histograms -- one corresponds to column `x`, one corresponds to column `y`, and one does not correspond to either column. 

**Histogram A:**  
![](var3.png)

**Histogram B:**  
![](var1.png)

**Histogram C:**  
![](var2.png)

#### Part 4.1 (5 pts)


 Suppose we run `t.hist('x')`. Which histogram does this code produce? Assign `histogram_column_x` to either 1, 2, or 3. 

1. Histogram A
2. Histogram B
3. Histogram C



```python
histogram_column_x = ...
```

```python
grader.check("q4.1")
```

<!-- BEGIN QUESTION -->

#### Part 4.2 (5 pts)


 State at least one reason why you chose the histogram from Part 4.1. **Make sure to indicate which histogram you selected** (ex: "I chose histogram A because ..."). 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Part 4.3 (5 pts)


 Suppose we run `t.hist('y')`. Which histogram does this code produce? Assign `histogram_column_y` to either 1, 2, or 3. 

1. Histogram A
2. Histogram B
3. Histogram C



```python
histogram_column_y = ...
```

```python
grader.check("q4.3")
```

<!-- BEGIN QUESTION -->

#### Part 4.4 (5 pts)


State at least one reason why you chose the histogram from Question 3.  **Make sure to indicate which histogram you selected** (ex: "I chose histogram A because ..."). 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

You're done with Lab 3!  

**Important submission information:** Follow these steps to submit your work:
* Run the tests and verify that they pass as you expect. 
* Choose **Save Notebook** from the **File** menu.
* **Run the final cell** and click the link below to download the zip file. 

Once you have downloaded that file, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to the corresponding assignment. The name of this assignment is "Lab 3 Autograder". **Be sure your work is saved before running the last cell!**

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1.1 (points: 5)
```
>>> unemployment.select('Date', 'NEI', 'NEI-PTER').take(0).column(0).item(0) == '1994-01-01'
True
```


Tests for q1.2 (points: 5)
```
>>> by_nei.take(0).column(0).item(0) == '2009-10-01'
True
```
```
>>> by_nei_pter.take(0).column(0).item(0) == '2009-10-01'
True
```
```
>>> # Run this cell to check your by_nei table. You do not need to change the code.
>>> by_nei.show(5)
<IPython.core.display.HTML object>
```
```
>>> # Run this cell to check your by_nei_pter table. You do not need to change the code.
>>> by_nei_pter.show(5)
<IPython.core.display.HTML object>
```


Tests for q1.3 (points: 5)
```
>>> greatest_nei.num_rows == 10
True
```
```
>>> greatest_nei.take(0).column(0).item(0) == '2009-10-01'
True
```


Tests for q1.4 (points: 5)
```
>>> round(pter.item(6), 4) == 1.1282
True
```


Tests for q1.5 (points: 5)
```
>>> by_pter.take(0).column(0).item(0) == '2009-07-01'
True
```


Tests for q1.6 (points: 5)
```
>>> pter_over_time.take(0).column("Date").item(0) == '1994-01-01'
True
```


Tests for q1.7 (points: 5)
```
>>> highPTER == True or highPTER == False
True
```


Tests for q2.1 (points: 5)
```
>>> 0 < us_birth_rate < 1
True
```


Tests for q2.2 (points: 5)
```
>>> 0 <= movers.num_rows <= 52
True
```


Tests for q2.3 (points: 5)
```
>>> 5e5 < west_births < 1e6
True
```
```
>>> west_births == 979657
True
```


Tests for q2.4 (points: 5)
```
>>> 1 <= visualization <= 3
True
```


Tests for q2.6 (points: 5)
```
>>> assoc == True or assoc == False
True
```


Tests for q3.3 (points: 5)
```
>>> 1 <= unit_meaning <= 3
True
```


Tests for q3.4 (points: 5)
```
>>> 0 <= boston_under_15 <= 100
True
```
```
>>> 0 <= manila_under_15 <= 100
True
```


Tests for q3.5 (points: 5)
```
>>> 1 <= boston_median_bin <= 5
True
```


Tests for q4.1 (points: 5)
```
>>> 1 <= histogram_column_x <= 3
True
```


Tests for q4.3 (points: 5)
```
>>> 1 <= histogram_column_y <= 3
True
```
