```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

# Lab 9: Sample Sizes and Confidence Intervals

**Helpful Resource:**
- [Python Reference](http://data8.org/fa21/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 
* [Ch 13. Estimation](https://inferentialthinking.com/chapters/13/Estimation.html)
* [Ch 14. Why the Mean Matters](https://inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.  For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines make plots look nice and hide some messy Python warnings.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
warnings.simplefilter('ignore', np.VisibleDeprecationWarning)
```

## 1. Bounding the Tail of a Distribution (15 pts)


A community has an average age of 45 years with a standard deviation of 5 years.

In each part below, fill in the blank with a percent that makes the statement true **without further assumptions**, and explain your answer.

*Note:* No credit will be given for loose bounds such as "at least 0%" or "at most 100%". Give the best answer that is possible with the information given.

<!-- BEGIN QUESTION -->

#### Part 1.1 (5 pts)


 At least _______% of the people are between 25 and 65 years old. 



_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part 1.2 (5 pts)


 At most _______% of the people have ages that are not in the range 25 years to 65 years. 



_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Part 1.3 (5 pts)


 At most _______% of the people are more than 65 years old. 

*Hint:* If you're stuck, try thinking about what the distribution may look like in this case.



_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Sample Size and Confidence Level (25 pts)


A data science class wants to estimate the percent of TikTok users among students at Williams. To do this, they need to take a random sample of students. You can assume that their method of sampling is equivalent to drawing at random with replacement from students at the school.

#### Part 2.1 (5 pts)


 Before starting this exercise, please review [Section 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html#the-sample-size) of the textbook. Your work will go much faster that way. 

Assign `smallest` to the smallest number of students they should sample to ensure that a **95%** confidence interval for the parameter has a width of no more than 6% from left end to right end. 

*Note:* While the true smallest sample size would have to be an integer, please leave your answer in decimal format for the sake of our tests.



```python
smallest = ...
smallest
```

```python
grader.check("q2.1")
```

<!-- BEGIN QUESTION -->

#### Part 2.2 (5 pts)


 Suppose the data science class decides to construct a 90% confidence interval instead of a 95% confidence interval, but they still require that the width of the interval is no more than 6% from left end to right end. Will they need the same sample size as in 2.1? Pick the right answer and explain further without calculation. 

1. Yes, they must use the same sample size.
2. No, a smaller sample size will work.
3. No, they will need a bigger sample.



_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Part 2.3 (5 pts)


 The professor tells the class that a 90% confidence interval for the parameter is constructed exactly like a 95% confidence interval, except that you have to go only 1.65 SDs on either side of the estimate (+/- 1.65) instead of 2 SDs on either side (+/- 2). Assign `smallest_num` to the smallest number of students they should sample to ensure that a **90%** confidence interval for the parameter has a width of no more than 6% from left end to right end? 

*Note:* While the true smallest sample size would have to be an integer, please leave your answer in decimal format for the sake of our tests.



```python
smallest_num = ...
smallest_num
```

```python
grader.check("q2.3")
```

For this next exercise, please consult [Section 14.3.4](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html#the-standard-normal-cdf) of the textbook for similar examples.

The students are curious about how the professor came up with the value 1.65 in Question 2.3. She says she ran the following two code cells. The first one calls the `datascience` library function `plot_normal_cdf`, which displays the proportion that is at most the specified number of SDs above average under the normal curve plotted with standard units on the horizontal axis. You can find the documentation [here](http://data8.org/datascience/util.html#datascience.util.plot_normal_cdf).

*Note:* The acronym `cdf` stands for `cumulative distribution function`. It measures the proportion to the left of a specified point under a probability histogram.

```python
plot_normal_cdf(1.65)
```

To run the second cell, the professor had to first import a Python library for probability and statistics:

```python
from scipy import stats
```

Then she used the `norm.cdf` method in the library to find the gold proportion above.

```python
stats.norm.cdf(1.65)
```

<!-- BEGIN QUESTION -->

#### Part 2.4 (5 pts)


 This shows that the percentage in a normal distribution that is at most 1.65 SDs above average is about **95%**. Explain why 1.65 is the right number of SDs to use when constructing a **90%** confidence interval. 



_Type your answer here, replacing this text._

```python
grader.check("q2.4")
```

<!-- END QUESTION -->

#### Part 2.5 (5 pts)


 The cell above shows that the proportion that is at most 2.33 SDs above average in a normal distribution is 99%. Assign `option` to the right option to fill in the blank: 

If you start at the estimate and go 2.33 SDs on either side, then you will get a _______% confidence interval for the parameter.

1. 99.5
2. 99
3. 98.5
4. 98



```python
option = ...
option
```

```python
grader.check("q2.5")
```

## 3. Polling and the Normal Distribution (45 pts)



Michelle is a statistical consultant, and she works for a group that supports Proposition 68 (which would mandate labeling of all horizontal and vertical axes), called Yes on 68.  They want to know how many Californians will vote for the proposition.

Michelle polls a uniform random sample of all California voters, and she finds that 210 of the 400 sampled voters will vote in favor of the proposition. We have provided a table for you below which has 3 columns: the first two columns are identical to `sample`. The third column contains the proportion of total voters that chose each option.

```python
sample = Table().with_columns(
    "Vote",  make_array("Yes", "No"),
    "Count", make_array(210,   190))

sample_size = sum(sample.column("Count"))
sample_with_proportions = sample.with_column("Proportion", sample.column("Count") / sample_size)
sample_with_proportions
```

#### Part 3.1 (5 pts)


 Michelle wants to use 10,000 bootstrap resamples to compute a confidence interval for the proportion of all California voters who will vote Yes.  

Fill in the next cell to simulate an empirical distribution of Yes proportions. Use bootstrap resampling to simulate 10,000 election outcomes, and assign `resample_yes_proportions` to contain the Yes proportion of each bootstrap resample. Then, visualize `resample_yes_proportions` with a histogram. **You should see a bell shaped curve centered near the proportion of Yes in the original sample.** 

*Hint:* `sample_proportions` may be useful here!



```python
resample_yes_proportions = make_array()
for i in np.arange(10000):
    resample = ...
    resample_yes_proportions = ...
Table().with_column("Resample Yes proportion", resample_yes_proportions).hist(bins=np.arange(.2, .8, .01))
```

```python
grader.check("q3.1")
```

<!-- BEGIN QUESTION -->

#### Part 3.2 (5 pts)


 Why does the Central Limit Theorem (CLT) apply in this situation, and how does it explain the distribution we see above? 



_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Part 3.3 (5 pts)


In a population whose members are 0 and 1, there is a simple formula for the **standard deviation of that population**:

$$\texttt{standard deviation of population} = \sqrt{(\text{proportion of 0s}) \times (\text{proportion of 1s})}$$

(Figuring out this formula, starting from the definition of the standard deviation, is an fun exercise for those who enjoy algebra.)

 Using only the Central Limit Theorem and the numbers of Yes and No voters in our sample of 400, *algebraically* compute the predicted standard deviation of the `resample_yes_proportions` array. Assign this number to `approximate_sd`. **Do not access the data in `resample_yes_proportions` in any way.** 

Remember that the standard deviation of the sample means can be computed from the population SD and the size of the sample (the formula above might be helpful). If we do not know the population SD, we can use the sample SD as a reasonable approximation in its place. [This section](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html#the-sd-of-all-the-sample-means) of the textbook also may be helpful.



```python
approximate_sd = ...
approximate_sd
```

```python
grader.check("q3.3")
```

#### Part 3.4 (5 pts)


 Compute the standard deviation of the array `resample_yes_proportions`, which will act as an approximation to the true SD of the possible sample proportions. This will help verify whether your answer to question 3.3 is approximately correct. 



```python
exact_sd = ...
exact_sd
```

```python
grader.check("q3.4")
```

#### Part 3.5 (5 pts)


 **Again, without accessing `resample_yes_proportions` in any way**, compute an approximate 95% confidence interval for the proportion of Yes voters in California. 

The cell below draws your interval as a red bar below the histogram of `resample_yes_proportions`; use that to verify that your answer looks right.

*Hint:* How many SDs corresponds to 95% of the distribution promised by the CLT? Recall the discussion in the textbook [here](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).



```python
lower_limit = ...
upper_limit = ...
print('lower:', lower_limit, 'upper:', upper_limit)
```

```python
grader.check("q3.5")
```

Your confidence interval should overlap the number 0.5.  That means we can't be very sure whether Proposition 68 is winning, even though the sample Yes proportion is a bit above 0.5.

The Yes on 68 campaign really needs to know whether they're winning.  It's impossible to be absolutely sure without polling the whole population, but they'd be okay if the standard deviation of the sample mean were only 0.005.  They ask Michelle to run a new poll with a sample size that's large enough to achieve that.  (Polling is expensive, so the sample also shouldn't be bigger than necessary.)

Michelle consults Chapter 14 of your textbook.  Instead of making the conservative assumption that the population standard deviation is 0.5 (coding Yes voters as 1 and No voters as 0), she decides to assume that it's equal to the standard deviation of the sample,

$$\sqrt{(\text{Yes proportion in the sample}) \times (\text{No proportion in the sample})}.$$

Under that assumption, Michelle decides that a sample of 9,975 would suffice.

#### Part 3.6 (5 pts)


 Does Michelle's sample size achieve the desired standard deviation of sample means? What SD would you achieve with a smaller sample size? A higher sample size? To explore this, first compute the SD of sample means obtained by using Michelle's sample size. 



```python
estimated_population_sd = ...
michelle_sample_size = ...
michelle_sample_mean_sd = ...
print("With Michelle's sample size, you would predict a sample mean SD of %f." % michelle_sample_mean_sd)
```

```python
grader.check("q3.6")
```

#### Part 3.7 (5 pts)


 Next, compute the SD of sample means that you would get from a smaller sample size. Ideally, you should pick a number that is significantly smaller, but any sample size smaller than Michelle's will do. 



```python
smaller_sample_size = ...
smaller_sample_mean_sd = ...
print("With this smaller sample size, you would predict a sample mean SD of %f" % smaller_sample_mean_sd)
```

```python
grader.check("q3.7")
```

#### Part 3.8 (5 pts)


 Finally, compute the SD of sample means that you would get from a larger sample size. Here, a number that is significantly larger would make any difference more obvious, but any sample size larger than Michelle's will do. 




```python
larger_sample_size = ...
larger_sample_mean_sd = ...
print("With this larger sample size, you would predict a sample mean SD of %f" % larger_sample_mean_sd)
```

```python
grader.check("q3.8")
```

#### Part 3.9 (5 pts)


 Based off of this, was Michelle's sample size approximately the minimum sufficient sample, given her assumption that the sample SD is the same as the population SD? Assign `min_sufficient` to `True` if 9,975 was indeed approximately the minimum sufficient sample, and `False` if it wasn't. 



```python
min_sufficient = ...
min_sufficient
```

You're done with Lab 9!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Go to [Gradescope](https://www.gradescope.com/) and submit the zip file to the corresponding assignment. The name of this assignment is "Lab 9 Autograder". 

**Be sure your work is saved before running the last cell.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q2.1 (points: None)
```
>>> 0 <= smallest <= 10000
True
```


Tests for q2.3 (points: None)
```
>>> 0 <= smallest_num <= 10000
True
```


Tests for q2.4 (points: None)
```
>>> # Just run this cell, do not change it.
>>> 
>>> stats.norm.cdf(2.33)
0.99009692444083575
```


Tests for q2.5 (points: None)
```
>>> 0 <= option <= 4
True
```


Tests for q3.1 (points: None)
```
>>> 0.45 <= np.mean(resample_yes_proportions) <= 0.55
True
```


Tests for q3.3 (points: None)
```
>>> 0 <= approximate_sd <= 1
True
```


Tests for q3.4 (points: None)
```
>>> 0 <= exact_sd <= 1
True
```
```
>>> round(exact_sd, 5) == 0.025
False
```


Tests for q3.5 (points: None)
```
>>> 0 <= lower_limit <= upper_limit < 1
True
```
```
>>> # Run this cell to plot your confidence interval.
>>> Table().with_column("Resample Yes proportion", resample_yes_proportions).hist(bins=np.arange(.2, .8, .01))
>>> plt.plot(make_array(lower_limit, upper_limit), make_array(0, 0), c='r', lw=10);
<Figure size 432x288 with 1 Axes>
```


Tests for q3.6 (points: None)
```
>>> 0 <= michelle_sample_mean_sd <= 1
True
```


Tests for q3.7 (points: None)
```
>>> 0 <= smaller_sample_mean_sd
True
```


Tests for q3.8 (points: None)
```
>>> 0 <= larger_sample_mean_sd
True
```


Tests for q3.9 (points: None)
